# Assignment — Power Law

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
from zlib import adler32
import requests
import networkx as nx
from typing import List
```

The PDF of the Power law distribution is 

$$ p(x) = Cx^{-\alpha},$$ 

where $C$ is a normalization constant 

$$C = \frac{\alpha - 1}{x_{\text{min}}^{-\alpha + 1}}$$ 

and $\alpha>1$ is called an exponent of the distribution. 

```python
def power_law_pdf(x, alpha=3.5, x_min=1):
    C = (alpha - 1) / x_min ** (1 - alpha)
    return C * x ** (-alpha)
```

```python
x_space = np.linspace(1, 10, 100)
plt.plot(x_space, power_law_pdf(x_space))
plt.title('Power Law PDF');
```

Let us generate observations from a Power Law RV.

### Task 1. Power law CDF (0.6 point)

The first step is to derive CDF of Powel law: $F(x) = P(X \leq x)$

$$F(x) = 1 - \int_{x}^\infty p(t) dt.$$

Take the integral, derive CDF analytically and write a function `power_law_cdf` with parameters `x`, `alpha` and `x_min`.

<!-- grade_id: cell-f682e125051975d6, points: 0, grade: false, solution: true -->
```python
def power_law_cdf(x, alpha=3.5, x_min=1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-87d0428f480773c4, points: 0.6, grade: true, solution: false -->
```python
assert power_law_cdf(2, 2, 1) == 0.5
assert power_law_cdf(10, 2, 1) == 0.9
assert adler32(str(power_law_cdf(2, 3, 1)).encode()) == 32571595
```

```python
x_space = np.linspace(1, 10, 100)
plt.plot(x_space, power_law_cdf(x_space))
plt.title('Power Law CDF');
```

### Task 2. Power law PPF (0.6 point)

Let $X \sim \text{Power law}$. Next, define a random variable $R$, s.t. $R = F(X)$, so $R$ will be uniformly distributed on interval [0, 1] ([proof](https://en.wikipedia.org/wiki/Probability_integral_transform#Proof)). Good thing here is that we easily can generate uniformly distributed pseudorandom numbers and then transform them into Power Law. Let us find an expression for $x = F^{-1}(r)$, where $r$ is an observation from uniform distrubution on interval [0, 1]. 

Find an analytical form of $F^{-1}(r)$ and write a function `power_law_ppf` (percent point function, also known as a quantile) with parameters `r`, `alpha` and `x_min`. 

<!-- grade_id: cell-63a819b2bc8441f1, points: 0, grade: false, solution: true -->
```python
def power_law_ppf(r, alpha=3.5, x_min=1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0b447767612a3798, points: 0.6, grade: true, solution: false -->
```python
assert round(power_law_ppf(0.5, 2, 1), 2) == 2
assert round(power_law_ppf(0.9, 2, 1), 2) == 10
assert adler32(str(round(power_law_ppf(0.96, 3, 1), 2)).encode()) == 19792020
```

```python
x_space = np.linspace(0, 0.999, 100)
plt.plot(x_space, power_law_ppf(x_space))
plt.title('Power Law PPF');
```

Finally, we can generate observation from Power law distribution as follows:
1. Generate observation from uniform distribution on interval [0, 1]
2. Calculate PPF value of given observation

```python
def power_law_generate(n, alpha=3.5, x_min=1, random_seed=1):
    np.random.seed(random_seed)
    uni_sample = np.random.uniform(0, 0.999, n)
    return power_law_ppf(uni_sample, alpha, x_min)
```

Look at the histogram of the generated sample

```python
alpha = 3.5
x_min = 1
x_train = power_law_generate(1000, alpha, x_min)
x_space = np.linspace(1, 15, 100)
```

```python
hist, bin_edges = np.histogram(x_train, bins=200, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers[hist > 0], hist[hist > 0], s=10)
plt.plot(x_space, power_law_pdf(x_space, alpha, x_min), 
         label='Theoretical PDF', c='tab:orange')
plt.legend()
plt.xlim(1, 10)
plt.show()
```

And here is the same histogram in log-log scale

```python
hist, bin_edges = np.histogram(x_train, bins=200, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers, hist, s=10)
plt.plot(x_space, power_law_pdf(x_space, alpha, x_min), 
         label='Theoretical PDF', c='tab:orange')
plt.legend()
plt.xscale('log')
plt.yscale('log')
plt.xlim(1, 10)
plt.show()
```

### Task 3. Estimation of alpha with linear binning (0.6 point)

Given observations from the Power Law distribution, try to estimate $\alpha$. The easiest way is to draw an empirical PDF with linear binning in log-log scale and apply linear regression. By _linear binning_ we mean to keep a bin width is fixed.

Write a function `alpha_lin_bins` that takes a train set, number of linear bins and returns an estimated $\alpha$.

_Hints:_
* _Take log in both side of $p(x) = Cx^{-\alpha}$_
* _To calculate an empirical PDF, use `np.histogram(x_train, bins=bins, density=True)`_
* _To calculate pseudoinverse matrix, use `np.linalg.pinv`_
* _Also you can use fit-predict in `sklearn.linear.LinearRegression`_

<!-- grade_id: cell-c57a5c7c5834b80e, points: 0, grade: false, solution: true -->
```python
def alpha_lin_bins(x_train, bins):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b9cf21aabb190bd7, points: 0.6, grade: true, solution: false -->
```python
bins = 100
x_train = power_law_generate(1000, x_min=1, alpha=2)
assert np.abs(alpha_lin_bins(x_train, bins) - 2) < 0.7
x_train = power_law_generate(1000, x_min=1, alpha=3.5)
assert np.abs(alpha_lin_bins(x_train, bins) - 3.5) < 0.9
x_train = power_law_generate(1000, x_min=1, alpha=10)
assert np.abs(alpha_lin_bins(x_train, bins) - 10) < 2
```

```python
x_min = 1
alpha = 3.5
bins = 100

x_train = power_law_generate(1000, alpha, x_min)
hist, bin_edges = np.histogram(x_train, bins=bins, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers, hist, s=10)

hat_alpha = alpha_lin_bins(x_train, bins)
x_space = np.linspace(x_min, x_train.max(), 100)
plt.plot(x_space, power_law_pdf(x_space, hat_alpha, x_min), 
         label='Estimated PDF', c='tab:orange')
plt.legend()
plt.title('Truth alpha = {:.2f}, estimated alpha = {:.2f}'.format(alpha, hat_alpha))
plt.xscale('log')
plt.yscale('log')
plt.show()
```

### Task 4. Estimation of alpha with logarithmic binning (2 points)

As we see the estimation with linear binning is noticeably inaccurate. For logarithmic binning we let the bin sizes increase with the value, making sure that each bin has a comparable number of observations.

Write a function `alpha_log_bins` that takes a train set, number of log bins and returns an estimated $\alpha$.

*Hint: use `np.logspace(np.log10(x_min), np.log10(x_train.max()), bins)` to generate an increasing logarithmic sequence of bins*

<!-- grade_id: cell-bc1092822f8126d1, points: 0, grade: false, solution: true -->
```python
def alpha_log_bins(x_train, bins):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1bc5eba53dc73dbd, points: 2, grade: true, solution: false -->
```python
bins = 100
x_train = power_law_generate(1000, x_min=1, alpha=2)
assert np.abs(alpha_log_bins(x_train, bins) - 2) < 0.2
x_train = power_law_generate(1000, x_min=1, alpha=3.5)
assert np.abs(alpha_log_bins(x_train, bins) - 3.5) < 0.5
x_train = power_law_generate(1000, x_min=1, alpha=10)
assert np.abs(alpha_log_bins(x_train, bins) - 10) < 1.5
```

```python
x_min = 1
alpha = 3.5
bins = 100

x_train = power_law_generate(1000, alpha, x_min)
binning = np.logspace(np.log10(x_min), np.log10(x_train.max()), bins)
hist, bin_edges = np.histogram(x_train, bins=binning, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers, hist, s=10)

hat_alpha = alpha_log_bins(x_train, bins)
x_space = np.linspace(x_min, x_train.max(), 100)
plt.plot(x_space, power_law_pdf(x_space, hat_alpha, x_min), 
         label='Estimated PDF', c='tab:orange')
plt.legend()
plt.title('Truth alpha = {:.2f}, estimated alpha = {:.2f}'.format(alpha, hat_alpha))
plt.xscale('log')
plt.yscale('log')
plt.show()
```

### Task 5. Generation graphs with a given degree distribution (0.6 point)

In this task we will generate different types of graph from its degree distributions.

__Graph A__ should be described with the density `[0, 2, 10]`.

```python
dens_a = [0, 2, 10]
plt.bar(range(len(dens_a)), dens_a);
```

<!-- grade_id: cell-4308485bfb1bf30e, points: 0, grade: false, solution: true -->
```python
def generate_graph_a() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-226da2d721eaace4, points: 0.15, grade: true, solution: false -->
```python
assert nx.degree_histogram(generate_graph_a()) == dens_a
```

__Graph B__

```python
dens_b = [0, 0, 10]
plt.bar(range(len(dens_b)), dens_b);
```

<!-- grade_id: cell-55d7cfacd5c8d553, points: 0, grade: false, solution: true -->
```python
def generate_graph_b() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-11bb8205ad27326a, points: 0.15, grade: true, solution: false -->
```python
assert nx.degree_histogram(generate_graph_b()) == dens_b
```

__Graph C__

```python
dens_c = [0, 0, 0, 0, 5]
plt.bar(range(len(dens_c)), dens_c);
```

<!-- grade_id: cell-6ad0ea3f81976162, points: 0, grade: false, solution: true -->
```python
def generate_graph_c() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-508d2e23c067c4d2, points: 0.15, grade: true, solution: false -->
```python
assert nx.degree_histogram(generate_graph_c()) == dens_c
```

__Graph D__

```python
dens_d = [0, 5, 0, 0, 0, 1]
plt.bar(range(len(dens_d)), dens_d);
```

<!-- grade_id: cell-2a6c0c0fe4dda491, points: 0, grade: false, solution: true -->
```python
def generate_graph_d() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c9c34144467f9031, points: 0.15, grade: true, solution: false -->
```python
assert nx.degree_histogram(generate_graph_d()) == dens_d
```

### Task 6. Reading a degree distribution from a file (2 points)

In this task we will estimate CDF of degree distribution of a given graph. Firstly, let us download the Game of Thrones relationships dataset and create a graph.

```python
raw_data = pd.read_csv("https://raw.githubusercontent.com/network-science-course/network-science-course/main/datasets/game_of_thrones_relationships.csv")

```

```python
raw_data.head()
```

```python
g = nx.from_pandas_edgelist(raw_data, source="character1", target="character2")
```

```python
nx.draw(g, node_size=10)
```

In the social network analysis, we will often consider the largest connected component only. Before estimation of CDF, we need to remove all connected components except of the largest one.

_Hint: use `nx.connected_components` to get a list of nodes for each connected component._

<!-- grade_id: cell-37477b43188105f5, points: 0, grade: false, solution: true -->
```python
def select_gygantic_component(g: nx.Graph) -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-00509dc9b42f24ea, points: 1, grade: true, solution: false -->
```python
gg = select_gygantic_component(g)
assert gg.number_of_edges() == 49
assert gg.number_of_nodes() == 37
```

```python
nx.draw(gg, node_size=30)
```

Now we can calculate an empirical CDF for a given graph. Recall a formal definition of CDF

$$F_X(x) = P(X \leq x)$$

And in empirical CDF, we estimate probability by share of nodes.

Write a method `empirical_cdf` that takes a graph and returns a np.array of probabilities. The first element is related to the node degree 0, the second — the node degree 1 and so on. The method should be able to group degrees occurence into equaly distributed bins. 

_Hint: use `nx.degee_histogram` to return a degree histogram — a list of numbers of nodes for each degree._

```python
nx.degree_histogram(g)
```

<!-- grade_id: cell-cbb478364b51233f, points: 0, grade: false, solution: true -->
```python
def empirical_cdf(g: nx.Graph) -> List[float]:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a1fd6b06aff25d78, points: 1, grade: true, solution: false -->
```python
ecdf = empirical_cdf(gg)
assert ecdf[0] == 0
assert ecdf[-1] == 1
assert np.all(ecdf[:-1] <= ecdf[1:])
assert adler32(str(empirical_cdf(gg).round(4).sum()).encode()) == 70123823
```

```python
plt.plot(ecdf)
plt.xlabel('degree')
plt.ylabel('probability')
plt.title('Empirical CDF')
plt.show()
```

### Task 7. MLE of alpha in Power Law (3 points)

Let us estimate $\alpha$ and $x_\min$ of a real social network using maximum likelihood. Here is a real degree distribution of the network of emails between members of the research institution with 1005 nodes.

```python
url = 'https://raw.githubusercontent.com/network-science-course/network-science-course/main/datasets/email_degree_sequence.txt'
open('email_degree_sequence.txt', 'wb').write(requests.get(url).content);

```

```python
degree_sequence = np.loadtxt('email_degree_sequence.txt')
hist, bin_edges = np.histogram(degree_sequence, bins=1000, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers[hist > 0], hist[hist > 0], s=5)
plt.title('Email network degree distribution')
plt.xscale('log')
plt.yscale('log')
```

The MLE consists of:
1. Fix $x_\min$ as a minimal node degree (drop node degrees that less than $x_\min$)
2. Calculate $\alpha$ via maximum likelihood estimation using fixed $x_\min$
$$\alpha = 1 + n \left[\sum_i \log \frac{x_i}{x_\min} \right]^{-1}$$
3. Calculate Kolmogorov-Smirnov test
4. Fix $x_\min$ as the next node degree
5. Repeat 2-4 by scanning all possible $x_\min$ and find the best $\alpha$ and $x_\min$ with respect to Kolmogorov-Smirnov test

Write a function `ml_power_law_params` that takes a node degree sequence `degree_sequence` and returns a tuple of two values: the best $\alpha$ and $x_\min$.

_Hints:_
* _To calculate Kolmogorov-Smirnov distance use `scipy.stats.kstest` with a parameter `cdf=power_law_cdf`_

<!-- grade_id: cell-c0da4153dae3178a, points: 0, grade: false, solution: true -->
```python
def mle_power_law_params(degree_sequence):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-213992ed205e9591, points: 3, grade: true, solution: false -->
```python
alpha, x_min = 3.5, 3
hat_alpha, hat_x_min = mle_power_law_params(power_law_generate(1000, alpha, x_min))
assert np.abs(hat_alpha - alpha) < 0.05
assert np.abs(hat_x_min - x_min) < 0.001

alpha, x_min = 7, 1
hat_alpha, hat_x_min = mle_power_law_params(power_law_generate(1000, alpha, x_min))
assert np.abs(hat_alpha - alpha) < 0.1
assert np.abs(hat_x_min - x_min) < 0.0001

hat_alpha, hat_x_min = mle_power_law_params(degree_sequence)
assert np.abs(hat_alpha - 2.9) < 0.1
assert np.abs(hat_x_min - 42) < 1
```

```python
hist, bin_edges = np.histogram(degree_sequence, bins=1000, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers[hist > 0], hist[hist > 0], s=5)
plt.title('Email network degree distribution')

hat_alpha, hat_x_min = mle_power_law_params(degree_sequence)
x_space = np.linspace(hat_x_min, degree_sequence.max(), 100)
plt.plot(x_space, power_law_pdf(x_space, hat_alpha, hat_x_min), 
         label='Estimated PDF', c='tab:orange')
plt.xscale('log')
plt.yscale('log')
plt.ylim(0.001, 0.5);
```

### Task 8. QQ plot (0.6 point)

It might be useful to draw a quantile-quantile plot (QQ plot) to compare empirical and theoretical distributions by the eye. Your task here is to generate Power Law observations by given parameters and compare theoretical and empirical quantiles.

Write a function `qqplot_data` that takes $\alpha$, $x_\min$, number of observations $n$ that will be generated and returns a tuple with two np.arrays: theoretical quantiles and empirical quantiles. Let quantiles be calculated for probabilities 0, 0.01, 0.02, ..., 0.99.

_Hint: to calculate an empirical quantile, use `np.quantile`_

<!-- grade_id: cell-7536b41d94152420, points: 0, grade: false, solution: true -->
```python
def qqplot_data(alpha, x_min, n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-cee74bedb007dbfa, points: 0.6, grade: true, solution: false -->
```python
x_min = 1
alpha = 3.5

theor100, empir100 = qqplot_data(alpha, x_min, 100)
theor200, empir200 = qqplot_data(alpha, x_min, 200)
theor1000, empir1000 = qqplot_data(alpha, x_min, 1000)
assert np.linalg.norm(theor1000 - empir1000) < np.linalg.norm(theor200 - empir200)
assert np.linalg.norm(theor200 - empir200) < np.linalg.norm(theor100 - empir100)
```

```python
plt.figure(figsize=(12, 3))

plt.subplots_adjust(wspace=0.3)

plt.subplot(1, 3, 1)
plt.scatter(theor100, empir100, s=10)
plt.plot([1, 7], [1, 7], '--', c='tab:orange')
plt.title('QQ plot, N = 100')
plt.xlabel('Theoretical distribution')
plt.ylabel('Empirical distribution')
plt.axis('square')

plt.subplot(1, 3, 2)
plt.scatter(theor200, empir200, s=10)
plt.plot([1, 7], [1, 7], '--', c='tab:orange')
plt.title('QQ plot, N = 200')
plt.xlabel('Theoretical distribution')
plt.ylabel('Empirical distribution')
plt.axis('square')

plt.subplot(1, 3, 3)
plt.scatter(theor1000, empir1000, s=10)
plt.plot([1, 7], [1, 7], '--', c='tab:orange')
plt.title('QQ plot, N = 1000')
plt.xlabel('Theoretical distribution')
plt.ylabel('Empirical distribution')
plt.axis('square');
```

```python

```