# Assignment — Introduction to network science and power law

```python
from zlib import adler32
import matplotlib.pyplot as plt
from typing import List
import numpy as np
import pandas as pd
import networkx as nx
import requests
```

In this assignment, we will consider some basics of networkx

```python
assert nx.__version__ >= '2.5'
```

In case of an exception, install networkx 2.5 using `pip` as follows:
```
pip uninstall networkx
pip install networkx==2.5
```

### Task 1. Reading graph from a file (0 points)

```python
url = 'https://raw.githubusercontent.com/netspractice/network-science/main/datasets/the_lord_of_the_rings_adjlist.txt'
open('the_lord_of_the_rings_adjlist.txt', 'wb').write(requests.get(url).content)

url = 'https://raw.githubusercontent.com/netspractice/network-science/main/datasets/the_lord_of_the_rings_edgelist.txt'
open('the_lord_of_the_rings_edgelist.txt', 'wb').write(requests.get(url).content);

```

A common representation of a graph is an adjacency matrix, but many real-world graphs are sparse (many zeros in an adjacency matrix), so in most cases data is stored in cheap sparse formats
1. edge list
2. adjacency list

Edge list is a list of pairs of nodes — a target and destination of edges. Here is an edge list of an interaction network of films "The Lord of the Rings" (2001-2003).

```python
with open('the_lord_of_the_rings_edgelist.txt') as f:
    print(f.read()[:197])
```

Adjacency list is a list of lists of nodes where the first node is a root and others are leafs

```python
with open('the_lord_of_the_rings_adjlist.txt') as f:
    print(f.read()[:190])
```

For most cases we can use given functions in networkx:

* `nx.read_adjlist`
* `nx.read_edgelist`
* `nx.read_weighted_edgelist`

In the further task, read graphs `the_lord_of_the_rings_edgelist.txt`, `the_lord_of_the_rings_adjlist.txt` and combine them into a single graph such that the same nodes are merged in a single node preserving all adjacent edges.

*Hint: use `graph.add_edges_from`*

<!-- grade_id: cell-3d2796d8ac544282, points: 0, grade: false, solution: true -->
```python
def read_and_merge_graphs(edgelist_path, adjlist_path):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d59c4e439b3aa75d, points: 0.0, grade: true, solution: false -->
```python
G = read_and_merge_graphs("the_lord_of_the_rings_edgelist.txt", 
                          "the_lord_of_the_rings_adjlist.txt")
assert set(G.neighbors('SAM')) == {
    'BILBO',
    'ELROND',
    'FRODO',
    'GALADRIEL',
    'GANDALF',
    'GOLLUM',
    'GORBAG',
    'SHAGRAT',
    'SMEAGOL'}
assert nx.number_of_nodes(G) == 35
assert nx.number_of_edges(G) == 98
```

### Task 2. Directed graphs (0 points)

The direction of edges could be crucial feature for many tasks, e.g. link analysis. Directed edges are also known as links or arcs. In networkx, the type of graph is defined by a constructor: `nx.Graph` for undirected graph, `nx.DiGraph` for directed graph. A directed graph can have an asymmetric adjacency matrix.

```python
diG = nx.DiGraph()
diG.add_edges_from([[0, 1], [1, 2]])
nx.to_numpy_array(diG)
```

Write a function `convert_to_digraph` that takes an undirected graph and returns a new directed graph where the direction is defined by the rule: the node $i$ follows to $j$ if the length of `name` of $i$ is at least the length of `name` of $j$. Keep the ordering of nodes and the node attribute `name`.

<!-- grade_id: cell-4ff100c55840771c, points: 0, grade: false, solution: true -->
```python
def convert_to_digraph(G):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bfab7d9d16583e9c, points: 0.0, grade: true, solution: false -->
```python
newG = nx.convert_node_labels_to_integers(G, label_attribute='name')
diG = convert_to_digraph(newG)
assert list(diG.nodes) == list(newG.nodes)
for i, j in diG.edges:
    assert len(diG.nodes[i]['name']) >= len(diG.nodes[j]['name'])
assert len(diG.edges) == 117
```

### Task 3. Basic statistics (0 points)

Networkx provides a simple interface for getting basic graph statistics (look at the documentation to recall definitions):
* `nx.radius`
* `nx.diameter`
* `nx.average_shortest_path_length`
* `nx.average_clustering`

For example, here are statistics of a path graph with 3 nodes:

```python
G = nx.path_graph(3)
print('Radius: ', nx.radius(G))
print('Diameter: ', nx.diameter(G))
print('Average shortest path length: ', nx.average_shortest_path_length(G))
print('Average clustering coefficient: ', nx.average_clustering(G))
nx.draw_networkx(G)
```

Your task is to create graphs by given statistics. We can create a graph using simple networkx interfaces. For example, we can create a graph by passing an edgelist.

```python
G = nx.Graph([[0, 1], [1, 2], [2, 0]])
nx.draw_networkx(G)
```

Or create an empty graph and then add edges one by one.

```python
G = nx.Graph()
G.add_edge(0, 1)
G.add_edge(1, 2)
G.add_edge(2, 0)
nx.draw_networkx(G)
```

Your task is to create graphs by given statistics.

Write a function `graph1` that returns a graph with a radius 3 and diameter 3.

<!-- grade_id: cell-1b71ec27d5c70dc5, points: 0, grade: false, solution: true -->
```python
def graph1():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-02d925c056e1f040, points: 0.0, grade: true, solution: false -->
```python
G = graph1()
assert nx.radius(G) == nx.diameter(G) == 3
nx.draw_networkx(G)
```

Write a function `graph2` that returns a graph with a diameter of at least 5 average shortest path lengths.

<!-- grade_id: cell-e6b06880520f2cf7, points: 0, grade: false, solution: true -->
```python
def graph2():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a6c86b83e11199d8, points: 0.0, grade: true, solution: false -->
```python
G = graph2()
assert nx.diameter(G) >= 5 * nx.average_shortest_path_length(G)
nx.draw_networkx(G)
```

Write a function `graph3` that returns a graph with an average clustering coefficient between 0.1 and 0.11.

<!-- grade_id: cell-eda0d295af4fc263, points: 0, grade: false, solution: true -->
```python
def graph3():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d4aeb27141847c83, points: 0.0, grade: true, solution: false -->
```python
G = graph3()
assert 0.1 <= nx.average_clustering(G) <= 0.11
nx.draw_networkx(G)
```

### Task 4. Graph layout (0 points)

A graph layout is a dictionary with a location of every node: keys are nodes and values are x, y positions. By default in `nx.draw_networkx`, it is the spring layout — a force-directed representation of the network treating edges as springs holding nodes close, while treating nodes as repelling objects.

```python
G = nx.barabasi_albert_graph(10, 2)
layout = nx.spring_layout(G)
layout
```

```python
nx.draw_networkx(G, pos=layout)
```

The Kamada Kawai layout is based on path-length cost-function. Sometimes, it is quite similar to the spring layout, however a significant difference is that the spring layout is random and the Kamada Kawai is deterministic.

```python
layout = nx.kamada_kawai_layout(G)
nx.draw_networkx(G, pos=layout)
```

Circular layout allows to draw a graph where nodes are in a circle.

```python
layout = nx.circular_layout(G)
plt.figure(figsize=(5, 5))
nx.draw_networkx(G, pos=layout)
```

Bipartite layout allows to draw bipartite graph.

```python
B = nx.Graph()
B.add_nodes_from([1, 2, 3, 4], bipartite=0)
B.add_nodes_from(["a", "b", "c"], bipartite=1)
B.add_edges_from([(1, "a"), (1, "b"), (2, "b"), (2, "c"), (3, "c"), (4, "a")])
layout = nx.bipartite_layout(B, [1, 2, 3, 4])
nx.draw_networkx(B, pos=layout)
```

There are other layouts in networkx, also there are some useful parameters that you can find in documentation. Next, let us visualize a coappearance network of characters in The Matrix — a 1999 American science fiction action film written and directed by the Wachowskis. The dataset is taken from https://moviegalaxies.com/.

```python
url = 'https://raw.githubusercontent.com/network-science-course/network-science-course/main/datasets/the_matrix_graph.gexf'
open('the_matrix_graph.gexf', 'wb').write(requests.get(url).content);

```

```python
def get_matrix_graph():
    G = nx.read_gexf('the_matrix_graph.gexf')
    mapping = {}
    for node in G.nodes:
        mapping[node] = G.nodes[node]['label']
    G = nx.relabel_nodes(G, mapping)
    return G
matrixG = get_matrix_graph()
plt.figure(figsize=(13, 7))
layout = nx.kamada_kawai_layout(matrixG)
nx.draw_networkx(matrixG, pos=layout)
```

Your task is to draw the matrix graph using the kamada kawai layout with nonlinear modification to obtain a denser layout away from the NEO node. Let us define a proximity coefficient using Gaussian kernel

$$\text{coef}(i) = \exp\left(-\frac{\text{distance}^2}{2\sigma^2}\right)$$

where a _distance_ is the Euclidean distance from the node $i$ to the NEO node, $\sigma$ is the constant. Then the nodes should be moved towards the NEO node such that

$$\text{distance} = \text{coef}(i) \times \text{distance}$$

The simplest way to do it is to center the layout so that the NEO node is in x=0 y=0 and then to multiply all node positions by the proximity coefficient.

Write a function `dense_neo_layout` that takes a layout, sigma and returns a new (dense) layout.

_Hint: use `np.linalg.norm` to calculate Euclidean norm._

<!-- grade_id: cell-6dc95d64261dffa6, points: 0, grade: false, solution: true -->
```python
def dense_neo_layout(layout, sigma):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fea67b2cc10bdbd6, points: 0.0, grade: true, solution: false -->
```python
layout = nx.kamada_kawai_layout(matrixG)
layout = dense_neo_layout(layout, sigma=1.5)
assert type(layout) == dict
assert np.all(layout['NEO'] == [0, 0])
assert np.all(layout['REX'].round(2) == [-0.67,  0.58])
assert adler32(str(layout['AGENT BROWN'].round(2)[0]).encode()) == 32309448
```

```python
plt.figure(figsize=(13, 7))
nx.draw_networkx(matrixG, layout)
```

### Task 5. Nodes visualization (2 points)

We can set a node size using `draw_networkx` with a parameter `node_size`. It should be a list of values ordered by `G.nodes`. For example, consider a simple path graph.

```python
G = nx.path_graph(3)
G.nodes
```

Let us set a size 200 for the node 0, 300 for 1, 400 for 2.

```python
layout = nx.kamada_kawai_layout(G)
nx.draw_networkx(G, pos=layout, node_size=[200, 300, 400])
```

To color nodes, we can use `node_color`. It should be a list of colors:
* a tuple of float RGB 0-1 values: (red, green, blue)
* a tuple of float RGBA 0-1 values: (red, green, blue, alpha)
* a string — color name

```python
nx.draw_networkx(
    G,
    pos=layout,
    node_color=[(0.4, 0.8, 0, 0.6), (0.7, 0.4, 1, 0.6), (1, 0, 0.5, 0.6)])
```

```python
nx.draw_networkx(
    G,
    pos=layout,
    node_color=['green', 'cyan', 'pink'])
```

Also we can use a colors from matplotlib colormaps. Discrete:

```python
colors = plt.cm.Set1.colors
nx.draw_networkx(
    G,
    pos=layout,
    node_color=[colors[0], colors[1], colors[2]])
```

Or continuous:

```python
rainbow = plt.cm.rainbow
nx.draw_networkx(
    G,
    pos=layout,
    node_color=[rainbow(0.2), rainbow(0.8), rainbow(0.9)])
```

Another option of coloring nodes is to pass a list of single values by the `node_color` and colormap by the `cmap` parameter. In this case, the lowest value will assigned to the first color in colormap and vise versa.

```python
nx.draw_networkx(
    G,
    pos=layout,
    cmap = plt.cm.hot,
    node_color=[1, 2, 3])
```

We can forcibly set a color range using `vmin` and `vmax`.

```python
nx.draw_networkx(
    G,
    pos=layout,
    cmap = plt.cm.hot,
    node_color=[1, 2, 3],
    vmin=0, vmax=4)
```

We can border nodes using `edgecolors` and `linewidths` parameters.

```python
nx.draw_networkx(
    G=G, 
    pos=layout,
    cmap = plt.cm.hot, 
    node_color=[1, 2, 3],
    vmin=0, vmax=4, 
    edgecolors='black', 
    linewidths=1)
```

Your task is to adjust a size and color of nodes in the matrix graph, size and color are taken from a node degree and node attribute `G.nodes[node]`.

```python
matrixG.nodes['NEO']
```

Write a function `drawing_nodes_params` that takes a graph and returns a dictionary with parameters for `nx.draw_networkx` of the form `{'G': graph, 'node_color': color, ...}`. Parameters should contain a graph, kamada kawai layout, node colors from `viz.color` attribute. The colors should be converted from RGB range 0-255 to RGB range 0-1. The node size is a node degree `G.degree[node]` multiplied by 100.

<!-- grade_id: cell-23d41af4fdb48d3c, points: 0, grade: false, solution: true -->
```python
def drawing_nodes_params(graph):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ce378d189964cb4f, points: 1.0, grade: true, solution: false -->
```python
matrixG = get_matrix_graph()
params = drawing_nodes_params(matrixG)
assert type(params) == dict
keys = ['G', 'pos', 'node_color', 'node_size']
assert len(set(params.keys()).intersection(keys)) == 4
assert np.all(params['node_color'][10].round(2) == [0.37, 0.79, 0.75])
assert adler32(str(params['node_color'][0][1].round(2)).encode()) == 19071123
assert params['node_size'][0] == 200
assert params['node_size'][1] == 600
assert nx.to_numpy_array(params['G']).sum() == 326.0
```

```python
plt.figure(figsize=(13, 7))
nx.draw_networkx(**params )
```

Sometimes, it will be useful to draw nodes in nonlinear scale. If we want to make node sizes more comparable, we can use polynomial or exponential scale. On the other hand, we can use logarithmic scale to make sizes more similar.

Write a function `drawing_nodes_polyparams` that returns the same dictionary as the `drawing_nodes_params` function, but the node size is a sum of node degree in a power 3 and 10.

$$\text{size} = \text{degree}^3 + 10$$

<!-- grade_id: cell-7056cc638b9de958, points: 0, grade: false, solution: true -->
```python
def drawing_nodes_polyparams(graph):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6b4ea9133d917597, points: 1.0, grade: true, solution: false -->
```python
matrixG = get_matrix_graph()
params = drawing_nodes_polyparams(matrixG)
assert type(params) == dict
keys = ['G', 'pos', 'node_color', 'node_size']
assert len(set(params.keys()).intersection(keys)) == 4
assert np.all(params['node_color'][10].round(2) == [0.37, 0.79, 0.75])
assert adler32(str(params['node_color'][0][1].round(2)).encode()) == 19071123
assert params['node_size'][0] == 18
assert params['node_size'][1] == 226
assert nx.to_numpy_array(params['G']).sum() == 326.0
```

```python
plt.figure(figsize=(13, 7))
nx.draw_networkx(**params )
```

### Task 6. Degree distribution (2 points)

In this task, we will try to guess a graph by its degree distribution.

__Graph A__ is described by the histogram `[0, 2, 10]` — 0 nodes with degree 0, 2 nodes with degree 1, 10 nodes with degree 2.

```python
dens_a = [0, 2, 10]
plt.bar(range(len(dens_a)), dens_a);
```

<!-- grade_id: cell-4308485bfb1bf30e, points: 0, grade: false, solution: true -->
```python
def generate_graph_a() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-226da2d721eaace4, points: 0.5, grade: true, solution: false -->
```python
G = generate_graph_a()
assert nx.degree_histogram(G) == dens_a
assert nx.is_connected(G)
nx.draw_networkx(G)
```

__Graph B__

```python
dens_b = [0, 0, 10]
plt.bar(range(len(dens_b)), dens_b);
```

<!-- grade_id: cell-55d7cfacd5c8d553, points: 0, grade: false, solution: true -->
```python
def generate_graph_b() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-11bb8205ad27326a, points: 0.5, grade: true, solution: false -->
```python
G = generate_graph_b()
assert nx.degree_histogram(G) == dens_b
assert nx.is_connected(G)
nx.draw_networkx(G)
```

__Graph C__

```python
dens_c = [0, 0, 0, 0, 5]
plt.bar(range(len(dens_c)), dens_c);
```

<!-- grade_id: cell-6ad0ea3f81976162, points: 0, grade: false, solution: true -->
```python
def generate_graph_c() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-508d2e23c067c4d2, points: 0.5, grade: true, solution: false -->
```python
G = generate_graph_c()
assert nx.degree_histogram(G) == dens_c
assert nx.is_connected(G)
nx.draw_networkx(G)
```

__Graph D__

```python
dens_d = [0, 5, 0, 0, 0, 1]
plt.bar(range(len(dens_d)), dens_d);
```

<!-- grade_id: cell-2a6c0c0fe4dda491, points: 0, grade: false, solution: true -->
```python
def generate_graph_d() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c9c34144467f9031, points: 0.5, grade: true, solution: false -->
```python
G = generate_graph_d()
assert nx.degree_histogram(G) == dens_d
assert nx.is_connected(G)
nx.draw_networkx(G)
```

### Task 7. Power law CDF (2 points)

Scale-free network has a Power law degree distribution. Let us take a closer look at this distribution. The PDF of the Power law is 

$$ p(x) = Cx^{-\alpha},$$ 

where $C$ is a normalization constant 

$$C = \frac{\alpha - 1}{x_{\text{min}}^{-\alpha + 1}}$$ 

and $\alpha>1$ is called an exponent of the distribution. 

```python
def power_law_pdf(x, alpha=3.5, x_min=1):
    C = (alpha - 1) / x_min ** (1 - alpha)
    return C * x ** (-alpha)
```

```python
x_space = np.linspace(1, 10, 100)
plt.plot(x_space, power_law_pdf(x_space))
plt.title('Power Law PDF');
```

Let us generate observations from a Power Law RV. The first step is to derive CDF of Powel law: $F(x) = P(X \leq x)$

$$F(x) = 1 - \int_{x}^\infty p(t) dt.$$

Take the integral, derive CDF analytically and then write a function `power_law_cdf` with parameters `x`, `alpha` and `x_min`.

<!-- grade_id: cell-f682e125051975d6, points: 0, grade: false, solution: true -->
```python
def power_law_cdf(x, alpha=3.5, x_min=1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-87d0428f480773c4, points: 2.0, grade: true, solution: false -->
```python
assert power_law_cdf(2, 2, 1) == 0.5
assert power_law_cdf(10, 2, 1) == 0.9
assert adler32(str(power_law_cdf(2, 3, 1)).encode()) == 32571595
```

```python
x_space = np.linspace(1, 10, 100)
plt.plot(x_space, power_law_cdf(x_space))
plt.title('Power Law CDF');
```

### Task 8. Power law PPF (2 points)

Let $X \sim \text{Power law}$. Next, define a random variable $R$, s.t. $R = F(X)$, so $R$ will be uniformly distributed on interval [0, 1] ([proof](https://en.wikipedia.org/wiki/Probability_integral_transform#Proof)). Good thing here is that we easily can generate uniformly distributed pseudorandom numbers and then transform them into Power Law. Let us find an expression for $x = F^{-1}(r)$, where $r$ is an observation from uniform distrubution on interval [0, 1]. 

Find an analytical form of $F^{-1}(r)$ and write a function `power_law_ppf` (percent point function, also known as a quantile) with parameters `r`, `alpha` and `x_min`. 

<!-- grade_id: cell-63a819b2bc8441f1, points: 0, grade: false, solution: true -->
```python
def power_law_ppf(r, alpha=3.5, x_min=1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0b447767612a3798, points: 2.0, grade: true, solution: false -->
```python
assert round(power_law_ppf(0.5, 2, 1), 2) == 2
assert round(power_law_ppf(0.9, 2, 1), 2) == 10
assert adler32(str(round(power_law_ppf(0.96, 3, 1), 2)).encode()) == 19792020
```

```python
x_space = np.linspace(0, 0.999, 100)
plt.plot(x_space, power_law_ppf(x_space))
plt.title('Power Law PPF');
```

Finally, we can generate observation from Power law distribution as follows:
1. Generate observation from uniform distribution on interval [0, 1]
2. Calculate PPF value of given observation

```python
def power_law_generate(n, alpha=3.5, x_min=1, random_seed=1):
    np.random.seed(random_seed)
    uni_sample = np.random.uniform(0, 0.999, n)
    return power_law_ppf(uni_sample, alpha, x_min)
```

Look at the histogram of the generated sample

```python
alpha = 3.5
x_min = 1
x_train = power_law_generate(1000, alpha, x_min)
x_space = np.linspace(1, 15, 100)
```

```python
hist, bin_edges = np.histogram(x_train, bins=200, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers[hist > 0], hist[hist > 0], s=10)
plt.plot(x_space, power_law_pdf(x_space, alpha, x_min), 
         label='Theoretical PDF', c='tab:orange')
plt.legend()
plt.xlim(1, 10)
plt.show()
```

And here is the same histogram in log-log scale

```python
hist, bin_edges = np.histogram(x_train, bins=200, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers, hist, s=10)
plt.plot(x_space, power_law_pdf(x_space, alpha, x_min), 
         label='Theoretical PDF', c='tab:orange')
plt.legend()
plt.xscale('log')
plt.yscale('log')
plt.xlim(1, 10)
plt.show()
```

### Task 9. Estimation of alpha with linear binning (2 points)

Given observations from the Power Law distribution, try to estimate $\alpha$. The easiest way is to draw an empirical PDF with linear binning in log-log scale and apply linear regression. By _linear binning_ we mean to keep a bin width is fixed.

Write a function `alpha_lin_bins` that takes a train set, number of linear bins and returns an estimated $\alpha$.

_Hints:_
* _Take log in both side of $p(x) = Cx^{-\alpha}$_
* _To calculate an empirical PDF, use `np.histogram(x_train, bins=bins, density=True)`_
* _To calculate pseudoinverse matrix, use `np.linalg.pinv`_
* _Also you can use fit-predict in `sklearn.linear.LinearRegression`_

<!-- grade_id: cell-c57a5c7c5834b80e, points: 0, grade: false, solution: true -->
```python
def alpha_lin_bins(x_train, bins):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b9cf21aabb190bd7, points: 2.0, grade: true, solution: false -->
```python
bins = 100
x_train = power_law_generate(1000, x_min=1, alpha=2)
assert np.abs(alpha_lin_bins(x_train, bins) - 2) < 0.7
x_train = power_law_generate(1000, x_min=1, alpha=3.5)
assert np.abs(alpha_lin_bins(x_train, bins) - 3.5) < 0.9
x_train = power_law_generate(1000, x_min=1, alpha=10)
assert np.abs(alpha_lin_bins(x_train, bins) - 10) < 2
```

```python
x_min = 1
alpha = 3.5
bins = 100

x_train = power_law_generate(1000, alpha, x_min)
hist, bin_edges = np.histogram(x_train, bins=bins, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers, hist, s=10)

hat_alpha = alpha_lin_bins(x_train, bins)
x_space = np.linspace(x_min, x_train.max(), 100)
plt.plot(x_space, power_law_pdf(x_space, hat_alpha, x_min), 
         label='Estimated PDF', c='tab:orange')
plt.legend()
plt.title('Truth alpha = {:.2f}, estimated alpha = {:.2f}'.format(alpha, hat_alpha))
plt.xscale('log')
plt.yscale('log')
plt.show()
```

```python

```